---
title: '2.4: Searching and Compilation Flags'
description: 'Here we distinguish between re.match() vs. re.search() and introduce compilation flags.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Lesson 2.4: Searching and Compilation Flags

```yaml
type: VideoExercise
key: 7c8a321651
xp: 50
```

`@projector_key`
d15d73f08145d2bf07a7ea82e5afa437
