---
title: 'Sample from Eric Wen'
description: 'Here we distinguish between re.match() vs. re.search() and introduce compilation flags.'
---

## Lesson 2.4: Searching and Compilation Flags

```yaml
type: VideoExercise
key: 7c8a321651
xp: 50
```

`@projector_key`
d15d73f08145d2bf07a7ea82e5afa437

---

## Insert exercise title here

```yaml
type: NormalExercise
key: c1c628bbb9
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
